---
title       : Probabilités conditionnelles
description : Nous allons dans ce chapitre introduire la notion très importante de conditionalité ainsi que le fameux théorème de Bayes.
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf


--- type:NormalExercise lang:r xp:100 skills:1 key:db32241f0e

## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
