---
title: 'Analyse combinatoire'
description: 'Nous allons dans ce chapitre introduire les opérateurs combinatoires utiles pour le calcul des probabilités.'
---

## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 7d083f35e1
lang: r
xp: 100
skills: 1
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
