---
title       : Espérance mathématique et variance
description : Nous allons dans ce chapitre introduire les notions d'espérance et de variance.
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf


---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 7d083f35e1
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```