---
title       : Analyse combinatoire
description : Nous allons dans ce chapitre introduire les opérateurs combinatoires utiles pour le calcul des probabilités.
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf

---


--- type:NormalExercise lang:r xp:100 skills:1 key:cbca316e2f
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
